# Project Euler: Problem 8

https://projecteuler.net/problem=8

The four adjacent digits in the 1000-digit number that have the greatest product are 9 × 9 × 8 × 9 = 5832.

(see the number below)

Find the thirteen adjacent digits in the 1000-digit number that have the greatest product. What is the value of this product?

Use NumPy for this computation

```python
import numpy as np
```

```python
d1000 = "\
73167176531330624919225119674426574742355349194934\
96983520312774506326239578318016984801869478851843\
85861560789112949495459501737958331952853208805511\
12540698747158523863050715693290963295227443043557\
66896648950445244523161731856403098711121722383113\
62229893423380308135336276614282806444486645238749\
30358907296290491560440772390713810515859307960866\
70172427121883998797908792274921901699720888093776\
65727333001053367881220235421809751254540594752243\
52584907711670556013604839586446706324415722155397\
53697817977846174064955149290862569321978468622482\
83972241375657056057490261407972968652414535100474\
82166370484403199890008895243450658541227588666881\
16427171479924442928230863465674813919123162824586\
17866458359124566529476545682848912883142607690042\
24219022671055626321111109370544217506941658960408\
07198403850962455444362981230987879927244284909188\
84580156166097919133875499200524063689912560717606\
05886116467109405077541002256983155200055935729725\
71636269561882670428252483600823257530420752963450\
"
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def strproduct(lst):
    return np.cumsum(lst)[-1]
bestprod=0
for n in np.arange(len(d1000)-14):
    print(n/(len(d1000)-12))
    num = np.array(int(d1000[int(n):int(n+13)]))
    prod = strproduct(num)
    if prod > bestprod:
        bestprod=prod
print(bestprod)

```

Output:
```
0.0
0.00101214574899
0.00202429149798
0.00303643724696
0.00404858299595
0.00506072874494
0.00607287449393
0.00708502024291
0.0080971659919
0.00910931174089
0.0101214574899
0.0111336032389
0.0121457489879
0.0131578947368
0.0141700404858
0.0151821862348
0.0161943319838
0.0172064777328
0.0182186234818
0.0192307692308
0.0202429149798
0.0212550607287
0.0222672064777
0.0232793522267
0.0242914979757
0.0253036437247
0.0263157894737
0.0273279352227
0.0283400809717
0.0293522267206
0.0303643724696
0.0313765182186
0.0323886639676
0.0334008097166
0.0344129554656
0.0354251012146
0.0364372469636
0.0374493927126
0.0384615384615
0.0394736842105
0.0404858299595
0.0414979757085
0.0425101214575
0.0435222672065
0.0445344129555
0.0455465587045
0.0465587044534
0.0475708502024
0.0485829959514
0.0495951417004
0.0506072874494
0.0516194331984
0.0526315789474
0.0536437246964
0.0546558704453
0.0556680161943
0.0566801619433
0.0576923076923
0.0587044534413
0.0597165991903
0.0607287449393
0.0617408906883
0.0627530364
```

<!-- grade_id: projecteuler8, points: 10, grade: true, solution: false -->
```python
assert True # leave this for grading
```